# Ejercicio de introducción a Python

# **Instrucciones**

Para completar este módulo de revisión, te recomendamos seguir estas instrucciones:

1. Completa las funciones proporcionadas en este notebook, pero **no** cambies el nombre de la función ni el nombre de los argumentos. Si lo haces, el autograder fallará y no recibirás ningún punto.
2. Ejecuta todas las celdas donde se definen funciones antes de ejecutar las celdas de prueba. ¡Las funciones deben existir antes de ser evaluadas!
3. Lee cuidadosamente los docstrings de las funciones. Contienen información adicional sobre cómo debe estructurarse el código (un [docstring](https://www.datacamp.com/community/tutorials/docstrings-python) es el texto que aparece entre triples comillas).
4. Algunas funciones pueden requerir múltiples salidas (los docstrings te indicarán cuáles). Asegúrate de devolverlas en el orden correcto.

## **¿Cómo funcionan las funciones?**

Es posible que aún no hayas visto funciones en Python, pero no te preocupes, las cubriremos en un caso futuro. Por ahora, lo principal que necesitas saber es:

1. Escribe tu código como lo harías normalmente, solo que esta vez dentro de la función justo debajo del comentario `### Your code here`.
2. Las funciones generalmente comienzan con la palabra clave `def` y terminan con `return some_variable`. No modifiques ninguna de estas estructuras. Si escribes tu código como se explicó anteriormente, todo estará bien.
3. El código dentro de la función debe estar indentado correctamente.
4. El docstring te indicará el nombre de la variable que debes devolver. La declaración `return` también proporciona esta información. Por ejemplo, si la función dice `return fancy_dict`, asegúrate de que la variable que defines como resultado también se llame `fancy_dict`.


## Ejercicio 1

A continuación, se describe un proceso burocrático ficticio. Transfórmalo en un bloque de declaraciones condicionales.

> Para calificar para los beneficios, debes tener al menos 18 años y ser ciudadano o residente permanente. Las personas que ganan más de 50,000 dólares al año no son elegibles.

No olvides que tu variable de resultado debe llamarse `eligible`.

## **Pistas:**

1. Si te resulta difícil seguir las condiciones, siempre puedes dibujar un árbol de decisión.
2. Para comprobar si al menos una de dos condiciones es verdadera, puedes usar lógica booleana.


<!-- grade_id: cell-ffff28d23d4fd7bb, points: 0, grade: false, solution: true -->
```python
def benefit_screening(applicant):
    """
    Determina si un solicitante es elegible para un beneficio gubernamental.
    
    Argumentos:
    `applicant`: Una lista de Python con toda la información relevante.
    
    Salida:
    `eligible`: Un valor booleano de Python.
    
    Ejemplo:
    Dado el argumento [22, False, True, 45000],
    la variable de salida `eligible` tendría un valor de `True`
    
    """
    
    age = applicant[0]
    citizen = applicant[1]
    resident = applicant[2]
    income = applicant[3]
       
    
    # TU CÓDIGO AQUÍ
    raise NotImplementedError() # Elimina esta línea cuando ingreses tu solución
    
    return eligible

```

Para probar tu función, puedes ejecutar este código (puedes modificar los valores en la lista a tu gusto):

```python
applicant_no_1 = [17,False,True,5000] # Age, Citizen, Resident, Income
benefit_screening(applicant_no_1)
```

## Ejercicio 2

Completa la siguiente función. Escribe las condiciones para que si la entrada `number` es par, la variable de salida `is_even` se establezca en `True`. Si la entrada no es par, `is_even` se establecerá en `False`.

## **Pistas:**

1. Un número es par si, al dividirlo por 2, el residuo es igual a 0.
2. Para encontrar el residuo de una división en Python, puedes usar el operador [**módulo**](https://www.freecodecamp.org/news/the-python-modulo-operator-what-does-the-symbol-mean-in-python-solved/). Así: `a % b`.


<!-- grade_id: cell-17df40f0291a6753, points: 0, grade: false, solution: true -->
```python
def is_this_number_even(number):
    """
    Determina si `number` es par.
    
    Argumentos:
    `number`: Un número entero.
    
    Salida:
    `is_even`: Un valor booleano de Python.
    
    Ejemplo:
    Dada la entrada 3, la salida debería ser `False`.
    """
    
    # TU CÓDIGO AQUÍ
    raise NotImplementedError() # Elimina esta línea cuando ingreses tu solución
    
    return is_even
```

## Ejercicio 3

Reescribe este bloque condicional utilizando declaraciones `elif`:

~~~python
if my_dict["color"] == "brown":
    out = "Chocolate"
else:
    if my_dict["color"] == "red":
        out = "Strawberries"
    else:
        if my_dict["color"] == "green":
            out = "Avocado"
        else:
            out = "Unknown food"        
~~~

<!-- grade_id: cell-fba7cf3b3a00d32d, points: 0, grade: false, solution: true -->
```python
def rewrite_conditional(my_dict):
    """
    Reescribe un bloque condicional.
    
    Argumentos:
    `my_dict`: Un diccionario de Python.
    
    Salida:
    `out`: Una cadena de texto en Python.
    
    Ejemplo:
    Dada la entrada {"color":"red"},
    la salida debería ser la cadena "Strawberries".
    """
    
    # TU CÓDIGO AQUÍ
    raise NotImplementedError() # Elimina esta línea cuando ingreses tu solución
    return out

```

## Celdas de Prueba

Ejecuta las siguientes celdas para verificar tus respuestas. Asegúrate de ejecutar primero las celdas donde escribiste tu solución antes de ejecutar las celdas de prueba, de lo contrario, obtendrás un `NameError` al comprobar tus respuestas.

<!-- grade_id: cell-e8b54fc508af9f32, points: 1, grade: true, solution: false -->
```python
# Ejercicio 1
applicant = [15, True, True, 1]
assert type(benefit_screening(applicant)) == type(True), "Ej. 1 - ¡Tu salida no parece ser un booleano! Puedes verificarlo con type()."
assert benefit_screening(applicant) == False, "Ej. 1 - Uno o más de tus condicionales parecen estar incorrectos. ¿Quizás tu condicional de edad?"
applicant = [20, False, False, 1]
assert benefit_screening(applicant) == False, "Ej. 1 - Uno o más de tus condicionales parecen estar incorrectos. ¿Quizás tu condicional de ciudadanía/residencia?"
applicant = [20, True, False, 1]
assert benefit_screening(applicant) == True, "Ej. 1 - Uno o más de tus condicionales parecen estar incorrectos. ¿Quizás tu condicional de ciudadanía/residencia?"
applicant = [20, False, True, 1]
assert benefit_screening(applicant) == True, "Ej. 1 - Uno o más de tus condicionales parecen estar incorrectos. ¿Quizás tu condicional de ciudadanía/residencia?"
applicant = [20, True, False, 10000000]
assert benefit_screening(applicant) == False, "Ej. 1 - Uno o más de tus condicionales parecen estar incorrectos. ¿Quizás tu condicional de ingresos?"
print("¡El ejercicio 1 parece correcto!")
```

<!-- grade_id: cell-39c8af6587e9ec0a, points: 1, grade: true, solution: false -->
```python
# Ejercicio 2
assert is_this_number_even(157) == False, "Ej. 2 - ¡Revisa tus condicionales! Probamos tu función con 157 y dijo que era par."
assert is_this_number_even(5) == False, "Ej. 2 - ¡Revisa tus condicionales! Probamos tu función con 5 y dijo que era par."
assert is_this_number_even(-8) == True, "Ej. 2 - ¡Revisa tus condicionales! Probamos tu función con -8 y dijo que era impar."
print("¡El ejercicio 2 parece correcto!")
```

<!-- grade_id: cell-78a27010bc41fdeb, points: 1, grade: true, solution: false -->
```python
# Ejercicio 3
import inspect
source_fun = inspect.getsource(rewrite_conditional)
assert "elif" in source_fun, "Ej. 3 - ¿Usaste declaraciones elif en tu función?"
d = {"color": "brown"}
assert rewrite_conditional(d) == "Chocolate", "Ej. 3 - Revisa tus condicionales. No están dando la salida correcta para 'brown'."
d = {"color": "red"}
assert rewrite_conditional(d) == "Strawberries", "Ej. 3 - Revisa tus condicionales. No están dando la salida correcta para 'red'."
d = {"color": "green"}
assert rewrite_conditional(d) == "Avocado", "Ej. 3 - Revisa tus condicionales. No están dando la salida correcta para 'green'."
d = {"color": "purple"}
assert rewrite_conditional(d) == "Unknown food", "Ej. 3 - Revisa tus condicionales. No están dando la salida correcta para 'purple'."
print("¡El ejercicio 3 parece correcto!")

```