# Ejercicios de introducción a python

A lo largo de esta introducción, completarás varias mini-tareas autoevaluadas diseñadas para probar tu comprensión de los casos anteriores.

**Instrucciones**

Los ejercicios requerirán que escribas fragmentos simples de código. Para completar esta parte, te recomendamos que sigas estas instrucciones:

1. Completa las funciones proporcionadas en este cuaderno, pero **no** cambies el nombre de la función ni el nombre(s) del(los) argumento(s). Si lo haces, el evaluador automático fallará y no recibirás ningún punto.
2. Ejecuta todas las celdas de definición de funciones antes de ejecutar las celdas de prueba. ¡Las funciones deben existir antes de ser evaluadas!
3. Lee atentamente las docstrings de las funciones. Contienen información adicional sobre cómo debería verse el código (una [docstring](https://www.datacamp.com/community/tutorials/docstrings-python) es el texto que aparece entre las comillas triples).
4. Algunas funciones pueden requerir varios resultados (las docstrings te dicen cuáles). Asegúrate de que se devuelvan en el orden correcto.

**¿Cómo funcionan las funciones?**

Ya viste algunas funciones de Python en el módulo de repaso anterior. Aquí tienes un recordatorio de algunas cosas que debes tener en cuenta:

1. Escribe tu código como lo harías normalmente, solo que esta vez escríbelo dentro de la función justo debajo del comentario `### TU CÓDIGO AQUÍ` comment.
2. Las funciones usualmente comienzan con la palabra clave `def` y terminan con `return` `some_variable`. No modifiques ninguna de estas. Si simplemente escribes tu código como explicamos anteriormente, estarás bien.
3. El código que escribas debe estar indentado.
4. La docstring te informa el nombre de la variable que tendrás que devolver. La declaración `return` te proporciona esta información también. Así que, si por ejemplo, la función dice `return fancy_dict`, debes asegurarte de que la variable que defines como tu resultado también se llame `fancy_dict.

Aquí tienes un par de preguntas de ejemplo ya resueltas:

### Ejemplo 1

Escribe un código que tome un valor numérico, lo multiplique por sí mismo y luego devuelva el resultado.

```python
def example_1(input_integer):
    """
    Esta es una docstring. Las docstrings están escritas en inglés simple
    (no son código) y están destinadas a ayudarte a recordar
    lo que hace tu función. En este caso, esta función
    asigna `input_integer` a una variable, lo multiplica por sí mismo,
    y luego devuelve el resultado. Toma un entero como entrada.

    Nota: De ahora en adelante, utilizaremos el símbolo ` para encerrar palabras
    que son nombres de variables, como acabamos de hacer con `input_integer`
    """
    result = input_integer * input_integer
    return result
```

Puedes verificar que tu función funciona llamándola de la siguiente manera:

```python
example_1(4)
```

Output:
```
16
```

### Ejemplo 2

Escribe un código que tome cualquier valor entero, lo sume a su propio valor dividido por 10, y devuelva el resultado.

```python
def example_2(input_number):
    """
    Toma un entero, súmalo a su propio valor dividido por 10,
    y devuelve el resultado.

    Argumentos: (Los argumentos son las entradas que pasas a tu función)
    `input_number`: Un entero
    """
    result = input_number + (input_number/10)
    return result
```

And checking:

```python
example_2(100)
```

Output:
```
110.0
```

¿Listo? ¡Aquí vienen las preguntas de codificación!

### Ejercicio 1

Escribe un código que encuentre el producto de 4 números y lo devuelva.

<!-- grade_id: cell-6da7fa29a659afd3, points: 0, grade: false, solution: true -->
```python
def multiply_these_numbers(a,b,c,d):
    """Multiplica `a`, `b`, `c` y `d`
    y devuelve el producto como la variable `result`.

    Argumentos: (Los argumentos son las entradas que pasas a tu función)
    `a`, `b`, `c` y `d`: Un entero
    """
    # TU CÓDIGO AQUÍ
    raise NotImplementedError() # Elimina esta línea cuando ingreses tu solución
    return result
```

### Ejercicio 2

Escribe un código que sume (concatene) 3 cadenas de texto diferentes y devuelva el resultado.

<!-- grade_id: cell-b74d0ff90247b40b, points: 0, grade: false, solution: true -->
```python
def add_these_strings(string1, string2, string3):
    """
    Concadena `string1`, `string2` y `string3`.

    Dado que los argumentos son cadenas de texto, deben estar rodeados
    por comillas cuando los pases como entradas. Así que, por ejemplo,
    llamarías (ejecutarías) esta función así:

    add_these_strings("texto uno", "texto dos", "texto tres")

    Tu función no debería agregar espacios innecesarios. La salida de la
    llamada anterior debería ser

    "texto unotexto dostexto tres"

    y NO

    "texto uno texto dos texto tres"
    """
    # TU CÓDIGO AQUÍ
    raise NotImplementedError() # Elimina esta línea cuando ingreses tu solución
    return result
```

### Ejercicio 3

Escribe un código que tome un número originalmente formateado como una cadena de texto, lo divida por la longitud de la cadena original y devuelva el resultado.

<!-- grade_id: cell-466de4b7eb3a9f70, points: 0, grade: false, solution: true -->
```python
def number_gymnastics(number_as_string):
    """
    Toma un número originalmente formateado como una cadena de texto,
    luego divídelo por la longitud de la cadena original.

    Argumentos:
    `number_as_string`: Un número, solo que está formateado como una cadena de texto.
    Por ejemplo: En lugar de pasar 15, pasas "15"
    (rodeado de comillas)

    Salida:
    `result`: Un número.

    Ejemplo:
    Si llamas a number_gymnastics("14"),
    la salida debería ser el número 7.0, porque
    14 dividido por 2 (la longitud de la cadena "14") es igual
    al número 7.0
    """
    # TU CÓDIGO AQUÍ
    raise NotImplementedError() # Elimina esta línea cuando ingreses tu solución
    return result
```

## Celdas de prueba

Ejecuta las celdas a continuación para verificar tus respuestas. Asegúrate de ejecutar primero tus celdas de solución antes de ejecutar las celdas a continuación, de lo contrario, recibirás un `NameError` al verificar tus respuestas.

<!-- grade_id: cell-e3ff00e471bec774, points: 1, grade: true, solution: false -->
```python
# Ex. 1
assert multiply_these_numbers(1,2,3,4) == 24, "Ej. 1 - Por favor, revisa tus cálculos! Recuerda que para multiplicar números en Python, usas el símbolo *"
assert multiply_these_numbers(2,3,4,5) == 120, "Ej. 1 - ¡Por favor, revisa tus cálculos! Recuerda que para multiplicar números en Python, usas el símbolo *"
print("Ejercicio 1 parece correcto!")
```

<!-- grade_id: cell-c0480287df6d8dd5, points: 1, grade: true, solution: false -->
```python
# Ex. 2
assert add_these_strings("data", "is", "cool") == "dataiscool", "Ej. 2 - Probamos tu función, y la salida no es del todo correcta aún. Recuerda que para concatenar dos cadenas de texto, puedes usar el símbolo +."
assert add_these_strings("I ", "need ", "space") == "I need space", "Ej. 2 - Probamos tu función y la salida aún no es del todo correcta. Recuerda que para concatenar dos cadenas de texto, puedes usar el símbolo +"
print("Ejercicio 2 parece correcto!")
```

<!-- grade_id: cell-bbb0af688a1c570a, points: 1, grade: true, solution: false -->
```python
# Ex. 3
assert number_gymnastics("270") == 90, "Ej. 3 - ¡Aún no está del todo correcto! Recuerda, primero transforma la entrada en cadena a un número (revisa la clase si necesitas un repaso sobre cómo hacerlo), y luego divídelo por la longitud de la cadena de entrada (nuevamente, consulta la clase para ver cómo encontrar la longitud)"
assert number_gymnastics("5") == 5, "Ej. 3 - ¡Aún no está del todo correcto! Recuerda, primero transforma la entrada en cadena a un número (revisa la clase si necesitas un repaso sobre cómo hacerlo), y luego divídelo por la longitud de la cadena de entrada (nuevamente, consulta la clase para ver cómo encontrar la longitud)"
print("Ejercicio 3 parece correcto!")
```