# Ejercicio de introducción a python

# **Instrucciones**

Para completar este módulo de revisión, te recomendamos seguir estas instrucciones:

1. Al escribir tu código, **no** utilices nombres de funciones diferentes a los proporcionados en los enunciados. Si lo haces, el autograder fallará y no recibirás puntos.
2. Ejecuta todas las celdas donde se definen funciones antes de ejecutar las celdas de prueba. ¡Las funciones deben existir antes de ser evaluadas!

## Ejercicio 1

Escribe una función llamada `cube_root` que tome un número no negativo como único argumento. Esta función debe calcular la raíz cúbica del número y devolverla con una declaración `return`.

**Nota:**
Para este ejercicio, asume que el usuario siempre ingresará un número no negativo; no es necesario incluir verificaciones de validez de entrada.

**Pista:**  
Recuerda que para $X \geq 0$,

$$
\sqrt[3]{X} = X^{1/3}
$$

<!-- grade_id: cell-27de682534bc2137, points: 0, grade: false, solution: true -->
```python
# Por favor, nombra tu función cube_root - Si usas otro nombre, la prueba fallará.
# Escribe tu función abajo (revisa el caso 0.5 para ver cómo definir una función en Python).
# TU CÓDIGO AQUÍ
raise NotImplementedError() # Elimina esta línea cuando ingreses tu solución
```

## Ejercicio 2

Escribe un docstring apropiado para esta función:

<!-- grade_id: cell-b75f1fddc7d0f9b1, points: 0, grade: false, solution: true -->
```python
def my_function(my_string, reps): # No cambies esta línea!
    
    # TU CÓDIGO AQUÍ
    raise NotImplementedError() # Elimina esta línea cuando ingreses tu solución
    
    my_string = my_string.upper()
    repeated = my_string * reps
        
    return repeated

```

## Ejercicio 3

Escribe una función llamada `process_number` que tome dos números enteros positivos como argumentos. Esta función debe:

1. Calcular la raíz cúbica del primer número entero (para esto, te recomendamos llamar a la función `cube_root` dentro de esta función).
2. Elevar esa raíz cúbica a la potencia cuyo exponente es el segundo número.

Por ejemplo, si los dos enteros son $x=27$ y $y=4$, entonces la salida de esta función debería ser:

$$
(\sqrt[3]{x})^y = (\sqrt[3]{27})^4 = 3^4 = 81
$$

### **Nota:**  
Para este ejercicio, asume que el usuario siempre ingresará números positivos; no es necesario incluir verificaciones de validez de entrada.


<!-- grade_id: cell-59e7e19785f29189, points: 0, grade: false, solution: true -->
```python
# Por favor, nombra tu función process_number - Si usas otro nombre, la prueba fallará.
# Escribe tu función abajo (revisa el caso 0.5 para ver cómo definir una función en Python).
# TU CÓDIGO AQUÍ
raise NotImplementedError() # Elimina esta línea cuando ingreses tu solución
```

## Exercise 4

Escribe una función llamada `split_number` que tome un número flotante como entrada y devuelva dos valores flotantes:

1. El primer flotante es la parte "entera" del número. Por ejemplo, si la entrada es $25.7$, este primer flotante debe ser $25.0$.
2. El segundo flotante es la parte "decimal" del número. Por ejemplo, si la entrada es $25.7$, este segundo flotante debe ser $7.0$ (*no* $0.7$).

**Pista:**
Para esto, te sugerimos que primero conviertas el número en una cadena y luego uses `.split(".")` para separar sus componentes. No olvides devolver ambos valores como **flotantes**.


<!-- grade_id: cell-e608ada7b2b43e72, points: 0, grade: false, solution: true -->
```python
def split_number(number):
    """
    Divide `number` en su parte entera y decimal.
    
    Argumentos:
    `number`: Un número entero o flotante.
    
    Salidas:
    `integer_part`: Un número flotante.
    `decimal_part`: Un número flotante.
    
    Ejemplo:
    split_number(25.7) debería producir 25.0 como primera salida
    y 7.0 como segunda salida.
    """
    # TU CÓDIGO AQUÍ
    raise NotImplementedError() # Elimina esta línea cuando ingreses tu solución
    return integer_part, decimal_part
```

## Celdas de Prueba

Ejecuta las siguientes celdas para verificar tus respuestas. Asegúrate de ejecutar primero las celdas donde escribiste tu solución antes de ejecutar las celdas de prueba, de lo contrario, obtendrás un `NameError` al comprobar tus respuestas.

<!-- grade_id: cell-293c1c13c0baad1c, points: 1, grade: true, solution: false -->
```python
# Ejercicio 1
try:
    cube_root(27)
except NameError:
    assert 3**3 == 9, "Ej. 1 - ¡No has definido una función llamada cube_root! Recuerda nombrar tu función cube_root y ejecutar la celda correspondiente antes de ejecutar esta prueba."
else:
    assert round(cube_root(27)) == 3, "Ej. 1 - ¡Incorrecto! Recuerda que para elevar un número x a la potencia y, se usa x**y en Python."
    assert round(cube_root(1000)) == 10, "Ej. 1 - ¡Incorrecto! Recuerda que para elevar un número x a la potencia y, se usa x**y en Python."
    print("¡El ejercicio 1 parece correcto!")

```

<!-- grade_id: cell-303c96e4e07385c0, points: 1, grade: true, solution: false -->
```python
# Ejercicio 2
import inspect
source_fun = inspect.getsource(my_function)
assert source_fun.count('"""') == 2, "Ej. 2 - ¿Agregaste un docstring a tu función? Recuerda usar comillas triples tanto al principio como al final del docstring."
assert len(source_fun.split('"""')[1]) > 15, "Ej. 2 - ¿Agregaste un docstring a tu función? Los docstrings cortos son buenos, pero a veces necesitamos más información."
assert "argument" in source_fun.split('"""')[1].lower(), "Ej. 2 - ¿Especificaste los argumentos de la función en tu docstring?"
assert "output" in source_fun.split('"""')[1].lower(), "Ej. 2 - ¿Especificaste la salida de la función en tu docstring?"
print("¡El ejercicio 2 parece correcto!")
```

<!-- grade_id: cell-0a378773c1097e13, points: 1, grade: true, solution: false -->
```python
# Ejercicio 3
try:
    process_number(27, 2)
except NameError:
    assert 3**3 == 9, "Ej. 3 - ¡No has definido una función llamada process_number! Recuerda nombrar tu función process_number y ejecutar la celda correspondiente antes de ejecutar esta prueba."
else:
    assert round(process_number(27, 2)) == 9, "Ej. 3 - Tu función no está produciendo la salida correcta. ¡Por favor, revisa tu aritmética!"
    assert round(process_number(64, 4)) == 256, "Ej. 3 - Tu función no está produciendo la salida correcta. ¡Por favor, revisa tu aritmética!"
    print("¡El ejercicio 3 parece correcto!")

```

<!-- grade_id: cell-9003ad0a93088895, points: 1, grade: true, solution: false -->
```python
# Ejercicio 4
assert split_number(27.2)[0] == 27, "Ej. 4 - Tu función no está produciendo la salida correcta. ¡Por favor, revisa tu aritmética!"
assert split_number(64.4)[1] == 4, "Ej. 4 - Tu función no está produciendo la salida correcta. ¡Por favor, revisa tu aritmética!"
assert split_number(12.0)[1] == 0, "Ej. 4 - Tu función no está produciendo la salida correcta. ¡Por favor, revisa tu aritmética!"
print("¡El ejercicio 4 parece correcto!")

```